<div class="admonition important alert alert-block alert-warning">
<p>⚠️ <strong>WICHTIG:</strong></p>
<p>
Die Auswertung der Programmieraufgaben erfolgt halbautomatisiert. Deshalb ist es wichtig, dass die Datei nur in den gekennzeichneten Bereichen verändert wird. Alle anderen Zellen sollen unberührt bleiben. Bitte keine Zellen löschen oder neue Zellen hinzufügen. Insbesondere die mit <code># Diese Zelle bitte NICHT verändern</code> markierten Zellen dienen der Auswertung und dürfen <strong>nicht</strong> verändert werden. Einige der darin befindlichen Tests sind aus dem Dokument entfernt, andere Testzellen wurden bewusst im Dokument gelassen und können gerne zur Prüfung des eigenen Codes ausgeführt (aber nicht verändert) werden.
</p></div>

```python
import numpy as np
import pyfar as pf
import scipy.signal as sig
import matplotlib.pyplot as plt
import ast
import inspect
%config InlineBackend.figure_format='retina'

def contains_for_loop(func):
    # This function checks if a function contains a for loop
    tree = ast.parse(func)
    for node in ast.walk(tree):
        if isinstance(node, ast.For):
            return True
    return False

print(f'pyfar version: {pf.__version__}')
version, minor_version, patch = pf.__version__.split('.')
if int(minor_version) < 6:
    print('If your version of pyfar does not support notch filters yet, update with "pip install pyfar --upgrade"')
```

# 3 Eigenschaften eines nichtlinearen Systems

# 3.1 $\mathrm{THD_R}$ (4P)
Schreiben Sie eine Funktion `thd_r(x, f_0)`, die `pf.Signal` Objekte annimmt und für jeden Kanal den auf den Effektivwert des **Gesamtsignals** bezogenen Klirrfaktor $\mathrm{THD_{R,\%}}$ in Prozent berechnet. Der auf das Gesamtsignal bezogene Klirrfaktor ist definiert als

$$ \text{THD}_\mathrm{R,\%} =  \frac{\sqrt{\sum_{k=2}^{n}U_{kf}^2}}{U_\text{ges}} \cdot 100\, \%.$$

$U_{kf}$: Effektive (rms) Spannung beim $k$-fachen der Anregefrequenz $f$ ($k$-te Harmonische) <br />
$U_\mathrm{ges}$: Effektive (rms) Spannung des Gesamtsignals

Dabei sollen alle Harmonischen benutzt werden, die kleiner als die Nyquist-Frequenz $f_\mathrm{s}/2$ sind. Mit der in der Funktion definierten Eingangsvariable `f_0` ist die Anregefrequenz $f$ gemeint.

Gehen Sie wie folgt vor:

1. Setzen Sie die `fft_norm` von `x` auf rms-Normalisierung.

2. Bestimmen Sie die Frequenzen aller Harmonischen, die kleiner als die Nyquistfrequenz $f_\mathrm{s}/2$ sind. Benutzen Sie z.B. `np.arange()`.

3. Bestimmen Sie die Indizes der Frequenzbins von `x`, die den Frequenzen der Harmonischen entsprechen. Benutzen Sie z.B. `x.find_nearest_frequency()`.

4. Berechnen Sie kanalweise die **Beträge** der rms-Amplituden an den Frequenzstützstellen, die den Harmonischen entsprechen (benutzen Sie `x.freq[...]` mit den korrekten Indizes). **Hinweis:** In `x.freq` ist das komplexe Spektrum mit Phaseninformationen gespeichert. Für die rms-Amplituden ist nur der Betrag relevant.

5. Bestimmen Sie kanalweise die rms-Amplitude des Gesamtsignals (**nicht** in dB). Benutzen Sie z.B. `pf.dsp.rms()`.

6. Berechnen Sie kanalweise den Klirrfaktor $\mathrm{THD_R}$ und speichern Sie die Werte in der Variable `thd_r`.

**Hinweis:** Benutzen Sie für die kanalweise Berechnung **keine** for-Schleifen. Die benötigten Pyfar-Funktionen arbeiten automatisch kanalweise. Die kanalweise Summierung der Spannungen der Harmonischen können Sie z.B. mit `np.sum(...,axis=...)` erledigen.

<!-- grade_id: t3-1-solution, points: 0, grade: false, solution: true -->
```python
def thd_r(x: pf.Signal, f_0: float):
    # fft norm von x auf 'rms' setzen:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Frequenzen aller Harmonischen, die kleiner als Nyquistfrequenz sind:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Indizes der Frequenzbins der Harmonischen:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # rms-Amplituden bei Harmonischen (NICHT in dB):
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Beträge der rms-Amplidude des Gesamtsignals (NICHT in dB):
    # (hier können Sie pf.dsp.rms() benutzen)
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # THD_R berechnen und in Variable thd_r speichern:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    if thd_r.shape == (1,):
        # Falls nur ein Wert berechnet wird, Skalarwert thd_r statt array([thd_r]) zurückgeben
        return thd_r.item()
    return thd_r  # thd_r Werte zurückgeben
```

<!-- grade_id: t-3-1-tests, points: 3, grade: true, solution: false -->
```python
# Diese Zelle bitte NICHT verändern
### BEGIN TESTS
assert not contains_for_loop(inspect.getsource(thd_r)), "The function contains a for loop, which is not allowed."

f_test = 10  # frequency
fs = 48000

# clean sine signal -> THD_R = 0
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=1, sampling_rate=fs)
assert(np.isclose(thd_r(x_test, f_test), 1.5e-15)), f'Expected: 1.5e-15, Got: {thd_r(x_test, f_test)}'

# clipped signal -> THD_R = 43%
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=1000, sampling_rate=fs)
x_test.time = np.clip(x_test.time, -1, 1)
assert(np.isclose(thd_r(x_test, f_test), 0.434848, atol=1e-6)), f'Expected: {0.434848 : %}, Got: {thd_r(x_test, f_test) : %}'

# Multi Channel (if this test fails, check that your calculations are done channel-wise)
levels = np.array([-90, 0, 90])
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=10**(levels/20), sampling_rate=fs)
noise = pf.signals.noise(n_samples=fs, rms=10**(-50/20), seed=1, sampling_rate=fs)
x_test += noise
x_test.time = np.clip(x_test.time, -1, 1)
assert(np.allclose(thd_r(x_test, f_test), [0.31463546, 0.00141489, 0.43484812], atol=1e-6)), f'Expected: [0.31463546, 0.00141489, 0.43484812], Got: {thd_r(x_test, f_test)}'
del f_test, fs, levels, noise, x_test
### END TESTS
```

# 3.2 $\mathrm{THD_F}$ (4P)
Schreiben Sie eine Funktion `thd_f(x, f_0)`, die `pf.Signal` Objekte annimmt und für jeden Kanal den auf die Anregefrequenz bezogenen Klirrfaktor $\mathrm{THD_{F,\%}}$ in Prozent berechnet. Der auf die Anregefrequenz bezogene Klirrfaktor ist definiert als

$$ \text{THD}_\mathrm{F,\%} =  \frac{\sqrt{\sum_{k=2}^{n}U_{kf}^2}}{U_{1f}} \cdot 100\, \%.$$

$U_{kf}$: Effektive (rms) Spannung beim $k$-fachen der Anregefrequenz $f$ ($k$-te Harmonische) <br />

Dabei sollen alle $n$ Harmonischen benutzt werden, die unterhalb der Nyquist-Frequenz $f_\mathrm{s}/2$ liegen.
Mit der in der Funktion definierten Eingangsvariable `f_0` ist die Anregefrequenz $f$ gemeint.

Gehen Sie wie folgt vor:

1. Setzen Sie die `fft_norm` von `x` auf rms-Normalisierung.

2. Bestimmen Sie die Frequenzen aller Harmonischen, die kleiner als die Nyquistfrequenz $f_\mathrm{s}/2$ sind. Benutzen Sie z.B. `np.arange()`.

3. Bestimmen Sie die Indizes der Frequenzbins von `x`, die den Frequenzen der Harmonischen entsprechen. Benutzen Sie z.B. `x.find_nearest_frequency()`.

4. Berechnen Sie kanalweise die **Beträge** der rms-Amplituden an den Frequenzstützstellen, die den Harmonischen entsprechen (benutzen Sie `x.freq[...]`  mit den korrekten Indizes). **Hinweis:** In `x.freq` ist das komplexe Spektrum mit Phaseninformationen gespeichert. Für die rms-Amplituden ist nur der Betrag relevant.

5. Berechnen Sie kanalweise den Betrag der rms-Amplitude an der Frequenzstützstelle, die der Fundamentalen entspricht (benutzen Sie `x.freq[...]`  mit den korrekten Indizes).

6. Berechnen Sie kanalweise den Klirrfaktor $\mathrm{THD_F}$ und speichern Sie die Werte in der Variable `thd_f`.

**Hinweis:** Benutzen Sie für die kanalweise Berechnung bitte **keine** for-Schleifen. Die benötigten Pyfar-Funktionen funktionieren automatisch kanalweise. Die kanalweise Summierung der Spannungen der Harmonischen können Sie z.B. mit `np.sum(...,axis=...)` erledigen.

<!-- grade_id: t-3-2-solution, points: 0, grade: false, solution: true -->
```python
def thd_f(x: pf.Signal, f_0: float):
    # fft norm von x auf 'rms' setzen:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Frequenzen aller harmonischen, die kleiner als Nyquistfrequenz sind:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Indizes der Frequenzbins der Harmonischen:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Beträge der rms-Amplituden der Harmonischen (NICHT in dB):
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Betrag der rms-Amplitude der Fundamentalen (NICHT in dB):
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # THD_F berechnen und in Variable thd_f speichern:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    if thd_f.shape == (1,):
        # Falls nur ein Wert berechnet wird, Skalarwert thd_f statt array([thd_f]) zurückgeben
        return thd_f.item()
    return thd_f  # thd_f Werte zurückgeben
```

<!-- grade_id: t-3-2-tests, points: 3, grade: true, solution: false -->
```python
# Diese Zelle bitte NICHT verändern
### BEGIN TESTS
assert not contains_for_loop(inspect.getsource(thd_f)), "The function contains a for loop, which is not allowed."

f_test = 10  # frequency
fs = 48000

# clean sine signal -> THD_f = 0
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=1, sampling_rate=fs)
assert(np.isclose(thd_f(x_test, f_test), 1.5e-15)), f'Expected: 1.5e-15, got: {thd_f(x_test, f_test)}'

# clipped signal -> THD_f = 48%
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=1000, sampling_rate=fs)
x_test.time = np.clip(x_test.time, -1, 1)
assert(np.isclose(thd_f(x_test, f_test), 0.482894, atol=1e-6)), f'Expected: {0.482894 : %}, got: {thd_f(x_test, f_test) : %}'

# Multi Channel (if only this test fails, check that your calculations are done channel-wise)
levels = np.array([-90, 0, 90])
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=10**(levels/20), sampling_rate=fs)
noise = pf.signals.noise(n_samples=fs, rms=10**(-50/20), seed=1, sampling_rate=fs)
x_test += noise
x_test.time = np.clip(x_test.time, -1, 1)
assert(np.allclose(thd_f(x_test, f_test), [32.408217,  0.001415,  0.482894], atol=1e-6)), f'Expected: [32.408217,  0.001415,  0.482894], got: {thd_f(x_test, f_test)}'
del f_test, fs, levels, noise, x_test
### END TESTS
```

# 3.3 $\text{THD+N}$ (3P)
Um nicht nur Verzerrungen bei ganzzahligen Vielfachen der Anregefrequenz, sondern auch alle anderen nichtlinearen Verzerrungen und Rauschen zu quantifizieren, soll nun eine Funktion `thd_n(x, f_0, Q=5)` geschrieben werden, welche `pf.Signal` Objekte annimmt und für jeden Kanal die Total Harmonic Distortion and Noise (THD+N) berechnet. THD+N ist nach [AES17](https://www.aes.org/publications/standards/search.cfm?docID=21) definiert als

$$ \text{THD+N}_\mathrm{\%} = \frac{U_{\text{Notch@}f_0}}{U_\text{ges}} $$

$U_{\text{Notch@}f_0}$: Effektive (**rms**) Spannung des Signals, bei dem die Anregefrequenz $f_0$ mit einem Notch-Filter herausgefiltert wurde <br />
$U_\mathrm{ges}$: Effektive (**rms**) Spannung des ungefilterten Gesamtsignals

Gehen Sie wie folgt vor:

1. Filtern Sie das Signal mit einem Notch-Filter `pf.dsp.filter.notch()` bei der Frequenz `f_0` mit Q-Faktor `Q`. Speichern Sie das gefilterte Signal in der Variable `x_filtered`.

2. Wenden Sie (wie in Aufgabe 1.15) Hann-Fensterung auf das gefilterte sowie auf das ungefilterte Signal an. Dabei soll die Fensterlänge genau der Signallänge entsprechen.

3. Berechnen Sie kanalweise die THD+N nach der oben definierten Formel. Speichern Sie das Ergebnis in der Variable `thd_n`. Hier dürfen Sie die Funktion `pf.dsp.rms()` benutzen.

**Hinweis:** Benutzen Sie für die kanalweise Berechnung **keine** for-Schleifen. Die benötigten Pyfar-Funktionen arbeiten automatisch kanalweise.

Wenn Ihre pyfar-Version keine Notch-Filter unterstützt, aktualisieren Sie das Paket (im Terminal oder Anaconda Prompt den Befehl `pip install pyfar --upgrade` eingeben). 

<!-- grade_id: t-3-3-solution, points: 0, grade: false, solution: true -->
```python
def thd_n(x: pf.Signal, f_0: float, Q=5):
    # Signal mit Notch-Filter filtern und in Variable x_filtered speichern:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # Hann-Fenster auf ungefiltertes und gefiltertes Signal anwenden:
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    # THD+N berechnen und in der Variable thd_n speichern::
    # ============= Ihre Lösung hier ===============
    raise NotImplementedError('Diese Zeile Löschen')
    # ==============================================

    if thd_n.shape == (1,):
        # Falls nur ein Wert berechnet wird, Skalarwert thd_n statt array([thd_n]) zurückgeben
        return thd_n.item()
    return thd_n  # thd_n Werte zurückgeben
```

<!-- grade_id: t-3-3-tests, points: 3, grade: true, solution: false -->
```python
# Diese Zelle bitte NICHT verändern
### BEGIN TESTS
assert not contains_for_loop(inspect.getsource(thd_n)), "The function contains a for loop, which is not allowed."

f_test = 1000  # frequency
fs = 48000

# clean sine signal -> THD+N = 0%
x_test = pf.signals.sine(frequency=f_test, n_samples=2*fs, amplitude=1, sampling_rate=fs)
assert(np.isclose(thd_n(x_test, f_test), 0, atol=1e-6)), f'Expected value near 0, Got: {thd_n(x_test, f_test)}'

# clipped signal -> THD+N = 39.5%
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=1000, sampling_rate=fs)
x_test.time = np.clip(x_test.time, -1, 1)
assert(np.isclose(thd_n(x_test, f_test), 0.395013, atol=1e-6)), f'Expected: {0.395013 : %}, Got: {thd_n(x_test, f_test) : %}'

# Multi Channel (if only this test fails, check that your calculations are done channel-wise)
levels = np.array([-90, 0, 90])
x_test = pf.signals.sine(frequency=f_test, n_samples=fs, amplitude=10**(levels/20), sampling_rate=fs)
noise = pf.signals.noise(n_samples=fs, rms=10**(-50/20), seed=1, sampling_rate=fs)
x_test += noise
x_test.time = np.clip(x_test.time, -1, 1)
assert(np.allclose(thd_n(x_test, f_test), [0.992724, 0.004358, 0.395014], atol=1e-6)), f'Expected: [0.992724, 0.004358, 0.395014], Got: {thd_n(x_test, f_test)}'
del f_test, fs, levels, noise, x_test
### END TESTS
```

# 3.4 $ \text{THD(+N)}$  vs Amplitude (3P)
Nun sollen die unterschiedlichen Berechnungsmethoden anhand von synthetischen Daten verglichen werden.
Dazu wurden Sinustöne bei der Frequenz `f_0` bei verschiedenen Amplituden erzeugt und dann mit einem nichtlinearen, digitalen System bearbeitet.

Die Datei `THD_vs_amplitude.far` enthält die Anregesignale `system_input`, die Systemantworten `system_output`, die Anregefrequenz `f_0`, die Amplituden `amplitudes` (digitale Amplitudenwerte ohne Einheit) und die entsprechenden Pegel `levels` in dBFS.

In den `pf.Signal`-Objekten `system_input` und `system_output` sind die Signale jeweils einer Amplitude in einem eigenen Kanal gespeichert. So enthält z.B. der Kanal `system_input[i]` einen Sinuston mit der in `amplitudes[i]` gespeicherten Amplitude, und der Kanal `system_output[i]` die entsprechende Systemantwort.

Gehen Sie wie folgt vor:

1. Laden Sie mit `pf.io.read()` die Datei `THD_vs_amplitude.far` in die Variable `measurement`. Die Variable enthält nun ein Python-Dictionary, das die oben genannten Objekte enthält. Geben Sie sich die Namen der Variablen mit `measurement.keys()` aus.

2. Speichern Sie die **effektiven** Anregepegel in der Variable `levels`, die Anregungsfrequenz in der Variable `f_0` und die Systemantworten in der Variable `system_output`. Verifizieren Sie, dass das Objekt `system_output` genau so viele Kanäle wie das Array `levels` Einträge hat.

3. Berechnen Sie die kanalweise $\mathrm{THD_{R, dB}}$, $\mathrm{THD_{F, dB}}$ und $\text{THD+N}_\mathrm{dB}$ von `system_output` (Umrechnung in dB siehe Einführungsveranstaltung) und speichern Sie die Werte in den Variablen `system_output_thd_r`, `system_output_thd_f` und `system_output_thd_n`.
   Benutzen Sie dazu Ihre vorher definierten Funktionen. 
   Verifizieren Sie, dass die Dimensionen der berechneten Werte den Dimensionen von `levels` entsprechen.
   Benutzen Sie für die Berechnung von $\text{THD+N}_\mathrm{dB}$ den Standardwert für den Gütefaktor `Q`.

4. Stellen Sie in einem gemeinsamen Plot $\mathrm{THD_{R, dB}}$, $\mathrm{THD_{F, dB}}$ und $\text{THD+N}_\mathrm{dB}$ der Systemantworten, jeweils aufgetragen über dem Anregepegel in dBFS, dar. Fügen Sie eine Legende hinzu und achten Sie auf korrekte Achsenbeschriftungen.

<!-- grade_id: t-3-4-solution, points: 2, grade: true, solution: true -->
```python
# 1.
# ============= Ihre Lösung hier ===============
raise NotImplementedError('Diese Zeile Löschen')
# ==============================================

# 2.
# ============= Ihre Lösung hier ===============
raise NotImplementedError('Diese Zeile Löschen')
# ==============================================

# 3.
# ============= Ihre Lösung hier ===============
raise NotImplementedError('Diese Zeile Löschen')
# ==============================================

# 4.
# ============= Ihre Lösung hier ===============
raise NotImplementedError('Diese Zeile Löschen')
# ==============================================
```

<!-- grade_id: t-3-4-tests, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte NICHT verändern
### BEGIN TESTS
assert system_output_thd_r.shape == system_output.cshape, f'Expected: {system_output.cshape}, Got: {system_output_thd_r.shape}'
assert system_output_thd_f.shape == system_output.cshape, f'Expected: {system_output.cshape}, Got: {system_output_thd_f.shape}'
assert system_output_thd_n.shape == system_output.cshape, f'Expected: {system_output.cshape}, Got: {system_output_thd_n.shape}'

assert np.allclose(system_output_thd_r[::5], [-24.473838,  -25.559581, -42.518391, -51.464693,  -9.452983],atol=1e-6)
assert np.allclose(system_output_thd_f[::5], [   8.331340, -17.522054, -42.431360, -51.464384,  -8.930134],atol=1e-6)
assert np.allclose(system_output_thd_n[::5], [  -0.167544,  -1.023732, -16.450333, -41.565004,  -9.471665],atol=1e-6)
### END TESTS
```

<!-- grade_id: t-3-4-2, points: 2, grade: true, solution: true -->
# 3.5 Verständnisfragen (5P)

Beschreiben und erklären Sie (kurz) die Verläufe von $\mathrm{THD_F}$, $\mathrm{THD_R}$ und $\text{THD+N}$. (2P)

---
_Ihre Lösung hier:_


---


<!-- grade_id: t-3-4-3, points: 1, grade: true, solution: true -->
Warum konnte der Fensterungs-Korrekturfaktor bei der Berechnung von $\text{THD+N}$ vernachlässigt werden? (1P)

---
_Ihre Lösung hier:_


---

<!-- grade_id: t-3-4-4, points: 1, grade: true, solution: true -->
Kann $\mathrm{THD_{R}}$ für ein Signal größer sein als $\text{THD+N}$? Begründen Sie Ihre Antwort. (1P)

---
_Ihre Lösung hier:_


---

<!-- grade_id: t-3-4-5, points: 1, grade: true, solution: true -->
Welche der Berechnungsverfahren können Werte größer als 100% (bzw. 0 dB) liefern? Begründen Sie Ihre Antwort. (1P)

---
_Ihre Lösung hier:_


---

# 3.6 Betriebskurve (4P)

Stellen Sie in einem Plot die rms-Pegel von `system_out` in dBFS über dem rms-Anregepegel in dBFS dar. Benutzen Sie die Funktion `pf.dsp.rms()`. Achten Sie auf korrekte Achsenbeschriftungen. (1P)

Stellen Sie in einem zweiten Plot die Zeitsignale aller Kanäle von `system_out` im Zeitintervall $[0\,\mathrm{ms}, 1\,\mathrm{ms}]$ dar. (1P)

**Tipp:** Die Zeiteinheit eines `pf.plot.time()`-Plots kann mit dem Parameter `unit` eingestellt werden.

<!-- grade_id: t-3-5-solution, points: 2, grade: true, solution: true -->
```python
# rms-Pegel vs Anregepegel:
# ============= Ihre Lösung hier ===============
raise NotImplementedError('Diese Zeile Löschen')
# ==============================================

# Zeitsignale:
# ============= Ihre Lösung hier ===============
raise NotImplementedError('Diese Zeile Löschen')
# ==============================================

```

<!-- grade_id: t-3-5-2-solution, points: 2, grade: true, solution: true -->
### Verständnisfragen (2P):
Beschreiben Sie die beiden Abbildungen. Welchen Pegel in dBFS hat das (unbewertete) Grundrauschen des Systems etwa? Um welche Art von Effekt könnte es sich bei dem zu untersuchenden System handeln? Begründen Sie Ihre Antwort.

---
_Ihre Lösung hier:_


---